Axonal distribution of mitochondria maintains neuronal autophagy during aging via eIF2β

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

The study presents significant findings on the role of mitochondrial depletion in axons and its impact on neuronal proteostasis. It effectively demonstrates how the loss of axonal mitochondria and elevated levels of eIF2β contribute to autophagy collapse and neuronal dysfunction. The use of Drosophila as a model organism and comprehensive proteome analysis adds robustness to the findings.

In this revision, the authors have responded thoughtfully to previous concerns. In particular, they have addressed the need for a quantitative analysis of age-dependent changes in eIF2β and eIF2α. By adding western blot data from multiple time points (7 to 63 days), they show that eIF2β levels gradually increase until middle age, then decline. In milton knockdown flies, this pattern appears shifted, supporting the idea that mitochondrial defects may accelerate aging-related molecular changes. These additions clarify the temporal dynamics of eIF2β and improve the overall interpretation.

Other updates include appropriate corrections to figures and quantification methods. The authors have also revised some of their earlier mechanistic claims, presenting a more cautious interpretation of their findings.

Overall, this work provides new insights into how mitochondrial transport defects may influence aging-related proteostasis through eIF2β. The manuscript is now more convincing, and the revisions address the main points raised earlier. I find the updated version much improved.

Thank you so much for the review, insightful comments and encouragement. We appreciate it.

Reviewer #2 (Public review):

In the manuscript, the authors aimed to elucidate the molecular mechanism that explains neurodegeneration caused by the depletion of axonal mitochondria. In Drosophila, starting with siRNA depletion of milton and Miro, the authors attempted to demonstrate that the depletion of axonal mitochondria induces the defect in autophagy. From proteome analyses, the authors hypothesized that autophagy is impacted by the abundance of eIF2β and the phosphorylation of eIF2α. The authors followed up the proteome analyses by testing the effects of eIF2β overexpression and depletion on autophagy. With the results from those experiments, the authors proposed a novel role of eIF2β in proteostasis that underlies neurodegeneration derived from the depletion of axonal mitochondria, which they suggest accelerates age-dependent changes rather than increasing their magnitude.

Strong caution is necessary regarding the interpretation of translational regulation resulting from the milton KD. The effect of milton KD on translation appears subtle, if present at all, in the puromycin incorporation experiments in both the initial and revised versions. Additionally, the polysome profiling data in the revised manuscript lack the clear resolution for ribosomal subunits, monosomes, and polysomes that is typically expected in publications.

Thank you so much for the review and insightful comments. We appreciate it.

Reviewer #2 (Recommendations for the authors):

The revised manuscript demonstrates many improvements. The authors have provided a more comprehensive data set and a more detailed description of their results. Furthermore, their explanation of the Integrated Stress Response (ISR) has been corrected, and this correction is reflected in the data interpretation.

As in the public review, I maintained my emphasis on the weakness of the claim on suppressed global translation, since the data are the same in the initial and the revised versions.

Thank you for your review. We understand that further studies will be needed to elucidate the roles on mitochondrial distribution in global translation profile. We will keep working on it.

A few suggestions for minor corrections.

(1) The order of figures in the revised version is disorganized.

Thank you for pointing it out. We corrected the order.

(2) In Figure 1A, mitochondria is bound by milton, and kinesin is bound by Miro. Their roles should be opposite.

Thank you for pointing it out, and we are sorry for the oversight. We corrected it.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Summary

Neuronal aging and neurodegenerative diseases are accompanied by proteostasis collapse, while the cellular factors that trigger it have not been identified. Impaired mitochondrial transport in the axon is another feature of aging and neurodegenerative diseases. Using Drosophila, we found that genetic depletion of axonal mitochondria causes dysregulation of protein degradation. Axons with mitochondrial depletion showed abnormal protein accumulation and autophagic defects. Lowering neuronal ATP levels by blocking glycolysis did not reduce autophagy, suggesting that autophagic defects are associated with mitochondrial distribution. We found that eIF2β was increased by the depletion of axonal mitochondria via proteome analysis. Phosphorylation of eIF2α, another subunit of eIF2, was lowered, and global translation was suppressed. Neuronal overexpression of eIF2β phenocopied the autophagic defects and neuronal dysfunctions, and lowering eIF2β expression rescued those perturbations caused by depletion of axonal mitochondria. These results indicate the mitochondria-eIF2β axis maintains proteostasis in the axon, of which disruption may underlie the onset and progression of age-related neurodegenerative diseases.

Introduction

Neurons have a morphologically complex architecture composed of microcompartments and require tight regulation of the abundance of proteins and organelles spatially and temporally1. Such control of protein amounts, or proteostasis, is essential for neuronal functions2 and is achieved through the orchestration of protein expression, folding, trafficking, and degradation controlled by intrinsic and environmental signals3. Translation is initiated by the eukaryotic initiation factor 2 (eIF2) complex4. eIF2, a heterotrimer of α, β, and γ subunits, transports Met-tRNA to the ribosome in a GTP-dependent manner5. Under stressed conditions, phosphorylation of eIF2α attenuates global translation and initiates translation of mRNAs related to the integrated stress response (ISR)6. As for protein degradation, autophagy and proteasome are major systems that maintain proteostasis7. The proteasome degrades unnecessary proteins, followed by regulated ubiquitination processes8, and autophagy removes damaged or harmful components, including large protein aggregates and organelles, through catabolism (selective autophagy)9. In addition to autophagy induced by acute stressors, a basal level of selective autophagy mediates the global turnover of damaged proteins10.

Such constitutive autophagy decreases during aging, which may underlie declines in the structural and functional integrity of neurons11. Decreased protein degradation and accumulation of abnormal proteins also contribute to increased risks of neurodegenerative diseases. Age-related neurodegenerative diseases such as Alzheimer’s disease and Parkinson’s disease are often associated with the accumulation of misfolded proteins such as amyloid-β, tau, and α-synuclein12. Enhancement of autophagy mitigates age-related dysfunctions and neurodegeneration caused by proteotoxic stress13. However, it is not fully understood how aging disrupts the regulation of this constitutive autophagy.

Neurons are also highly energy-demanding cells. At nerve terminals, action potentials trigger the release of neurotransmitters via exocytosis of synaptic vesicles, which requires a constant supply of ATP and calcium buffering14. Such neuronal activity relies on mitochondrial functions15, and mitochondria are actively transported from their major sites of biogenesis in soma to axons16. However, the axonal transport of mitochondria declines during aging17,18,19. Reduced axonal transport of mitochondria is thought to contribute to age-related declines in neuronal functions17,19,20,21. The number of functional mitochondria in synapses is reduced in the brains of patients suffering from age-related neurodegenerative diseases such as Alzheimer’s disease22, and mutations in genes involved in mitochondrial dynamics are linked to neurodegenerative diseases23. The mislocalization of mitochondria is sufficient to cause age-dependent neurodegeneration in Drosophila and mice24,25, indicating that the proper distribution of mitochondria is essential to maintain neuronal functions. Thus, depletion of functional mitochondria from axons and proteostasis collapse are common features of aging and neurodegenerative diseases.

Mitochondrial transport is regulated by a series of molecular adaptors that mediate the attachment of mitochondria to molecular motors16. In Drosophila, mitochondrial transport is mediated by milton and Miro, which attaches mitochondria to microtubules via kinesin heavy chain26,27. In the absence of milton or Miro, synaptic terminals and axons lack mitochondria, although mitochondria are numerous in the neuronal cell body28. We previously reported that RNAi-mediated knockdown of milton or Miro in neurons causes a reduction in axonal mitochondria, age-dependent locomotor defects29, and age-dependent neurodegeneration in neuropile area starting around 30 days after eclosion (day-old)24, and enhances axon degeneration caused by human tau proteins24, suggesting that these flies can be used as a model to analyze the effect of depletion of axonal mitochondria during aging. In this study, we investigated a causal relationship between mitochondrial distribution and neuronal proteostasis by using neuronal knockdown of milton. We found that depletion of axonal mitochondria reduced autophagy and increased the accumulation of aggregated proteins in the axon prior to gross neurodegeneration. Proteome analysis and follow-up biochemical analyses revealed that neuronal knockdown of milton increased eIF2β levels and lowered phosphorylation of eIF2αin the axon. In addition, milton knockdown suppressed global translation. Overexpression of eIF2β was sufficient to decrease autophagy and induce neuronal dysfunction, and genetic suppression of eIF2β restored autophagy and improved neuronal function in the milton knockdown background. These findings suggest that loss of axonal mitochondria and elevated levels of eIF2β mediate proteostasis collapse and neuronal dysfunction during aging.

Results

Depletion of axonal mitochondria by knockdown of milton or Miro causes protein accumulation in the axon

In Drosophila, mitochondrial transport is mediated by milton and Miro, which attach mitochondria to microtubules via kinesin heavy chain26,27 (Figure 1A). It has been reported that expression of milton RNAi in neurons via pan-neuronal elav-GAL4 driver reduced milton protein levels in Drosophila head lysate to 40% and mito-GFP signals in axons to 50%24,29.

Knockdown of milton or Miro causes protein accumulation in the axon

(A) Schematic representation of the mitochondrial transport machinery. Knockdown of milton, an adapter protein for mitochondrial transport, depletes mitochondria in the axon. (B, C) Ubiquitinated proteins in brains with neuronal knockdown of milton or Miro. Brains dissected at 14-day-old (B) or 30-day-old (C) were immunostained with an antibody against ubiquitin. Firefly luciferase RNAi was used as a control. Representative images (left) and quantitation of the number of ubiquitin-positive puncta (right) are shown. Scale bars of hemibrains, 100 µm, Scale bars of high magnifications, 10 µm. Means ± SE, n = 8. N.S., p > 0.05; ***p < 0.005 (one-way analysis of variance (ANOVA) followed by Tukey’s honestly significant difference (HSD) post hoc test). (D) Cross sections in the lamina and in the retina were used to analyze the ultrastructure of synapses and cell bodies, respectively. milton RNAi was expressed in the retina and neurons via a combination of GAL4 drivers, a pan-retinal gmr-GAL4 and pan-neuronal elav-GAL4. (E) Quantitation of the number of mitochondria in a presynaptic terminal from transmission electron micrographs. 180 presynaptic terminals from cross-sections of the lamina from three brains were analyzed. ***p < 0.005 (Chi-square test). (F, G) Presynaptic terminals of photoreceptor neurons of control and milton knockdown flies. Photoreceptor neurons are highlighted in blue. Swollen presynaptic terminals (asterisks in (F)), characterized by the enlargement and higher circularity, were found more frequently in milton knockdown neurons. Arrowheads indicate presynaptic terminals with dense materials. Scale bars, 2 µm. Representative images (Left) and quantitation (Right) are shown. 918–1118 from three heads were quantified for the percentage of swollen presynaptic terminals, and 180 presynaptic terminals from three heads were quantified for the size of presynaptic terminals. Mean ± SE, **p < 0.01, ***p < 0.005 (Student’s t-test). (G) Dense materials (arrowheads in (G)) in the presynaptic terminals of milton knockdown neurons. Scale bars, 2 µm. The ratio of presynaptic terminals containing dense materials were quantified from 918–1118 presynaptic terminals from three heads. Mean ± SE, ***p < 0.005 (Student’s t-test). (H) Cell bodies of photoreceptor neurons of control and milton knockdown flies. Scale bars, 2 µm. Flies were 27-day-old.

To test how loss of axonal mitochondria affects proteostasis in neurons, we first examined the accumulation of ubiquitinated proteins. At 14-day-old, more ubiquitinated proteins were deposited in the brains of milton knockdown flies than in those of age-matched control flies (Figure 1B, p < 0.005 between control RNAi and milton RNAi). There was no significant increase in ubiquitinated proteins in milton knockdown flies at 1-day old, suggesting that the accumulation of ubiquitinated proteins caused by milton knockdown is age-dependent (Figure S1). We also analyzed the effect of the neuronal knockdown of Miro, a partner of milton, on the accumulation of ubiquitin-positive proteins. Since severe knockdown of Miro in neurons causes lethality, we used UAS-Miro RNAi strain with low knockdown efficiency, whose expression driven by elav-GAL4 caused 30% reduction of Miro mRNA in head extract24. Although there was a tendency for increased ubiquitin-positive puncta in Miro knockdown brains, the difference was not significant (Figure 1B, p>0.05 between control RNAi and Miro RNAi). These data suggest that the depletion of axonal mitochondria induced by milton knockdown leads to the accumulation of ubiquitinated proteins before neurodegeneration occurs.

It has been reported that ubiquitinated proteins accumulate with aging30; thus, we analyzed the accumulation of ubiquitinated proteins in aged brains (30-day-old) with milton knockdown. The number of puncta of ubiquitinated proteins did not significantly differ between control and milton knockdown flies or between control and Miro knockdown flies (Figure 1C, p > 0.05). These results suggest that depletion of axonal mitochondria may have more impact on proteostasis in young neurons than in old neurons.

We examined the ultrastructure of presynaptic terminals and cell bodies in photoreceptor neurons with milton knockdown by transmission electron microscopy in 27-day-old flies (Figure 1D). As previously reported24, the number of mitochondria in presynaptic terminals decreased in milton knockdown (Figure 1E). The swelling of presynaptic terminals, characterized by the enlargement and roundness, was not reported at 3-day-old24 but observed at this age with about 4% of total presynaptic terminals (Figure 1F, asterisks).

Some presynaptic terminals of milton knockdown neurons contained dense materials (Figure 1F and G, arrowheads). Dense materials are rarely found in age-matched control neurons, indicating that milton knockdown induces abnormal protein accumulation in the presynaptic terminals (Figure 1G and H). In milton knockdown neurons, dense materials are found in swollen presynaptic terminals more often than in presynaptic terminals without swelling, suggesting a positive correlation between the disruption of proteostasis and axonal damage (Figure 1G). In contrast, dense materials were not observed in cell bodies in the milton knockdown retina (Figure 1H). These results indicate that the depletion of axonal mitochondria induces protein accumulation in the axon.

Depletion of axonal mitochondria impairs protein degradation pathways

Since abnormal proteins were accumulated in milton knockdown brains, we next examined if protein degradation pathways were suppressed. We analyzed autophagy via western blotting of the autophagy markers LC3 and p6231. During autophagy progression, LC3 is conjugated with phosphatidylethanolamine to form LC3-II, which localizes to isolation membranes and autophagosomes. LC3-I accumulation occurs when autophagosome formation is impaired, and LC3-II accumulation is associated with lysosomal defects31,32. p62 is an autophagy substrate, and its accumulation suggests autophagic defects31,32. We found that milton knockdown increased LC3-I, and the LC3-II/LC3-I ratio was lower in milton knockdown flies than in control flies at 14-day-old (Figure 2A). We also analyzed p62 levels in head lysates sequentially extracted using detergents with different stringencies (1% Triton X-100 and 2% SDS). Western blotting revealed that p62 levels were increased in the brains of 14-day-old of milton knockdown flies (Figure 2B). The increase in the p62 level was significant in the Triton X-100-soluble fraction but not in the SDS-soluble fraction (Figure 2B), suggesting that depletion of axonal mitochondria impairs the degradation of less-aggregated proteins. Proteasome activity was also significantly decreased in brains with neuronal knockdown of milton (Figure 2C, p < 0.005).

milton knockdown impairs protein degradation pathways

(A, B) Western blotting of head extracts of control and milton knockdown flies with antibodies against LC3 (A) and Ref2P, the fly homolog of mammalian p62 (B). For the analyses of p62 levels, heads were extracted with 1% Triton X-100 or 2% SDS (B). Flies were 14-day-old. Representative blots (left) and quantitation (right) are shown. Actin was used as a loading control. Means ± SE, n = 6 (LC3), n = 3 (p62). (C) Proteasome activity in head extracts of control and milton knockdown flies was measured by hydrolysis of Suc-LLVY-AMC at 14-day-old. Means ± SE, n = 3. (D, E) Western blotting of head extracts of 30-day-old control and milton knockdown flies. Blotting was performed with anti-LC3 (D) and anti-p62 (E) antibodies. Representative blots (left) and quantitation (right) are shown. Actin was used as a loading control. Means ± SE, n = 6 (LC3), n = 3 (p62). (F) Proteasome activity in head extracts of 30-day-old control and milton knockdown flies. Means ± SE, n = 3. N.S., p > 0.05; *p < 0.05; **p < 0.01; ***p < 0.005 (Student’s t-test).

At 30-day-old, LC3-I was still higher, and the LC3-II/LC3-I ratio was lower, in milton knockdown compared to the control (Figure 2D). At this age, milton knockdown increased p62 significantly in 1% Triton X-100 fraction and 2% SDS fraction (Figure 2E). Proteasome activities were also decreased in milton knockdown flies at 30-day-old (Figure 2F). These results indicate that depletion of axonal mitochondria impairs protein degradation pathways.

ATP deprivation does not impair autophagy

milton knockdown downregulates ATP in the axon33. To examine whether the disruption of protein degradation pathways by milton knockdown is due to ATP deprivation, we investigated the effects of knocking down phosphofructokinase (Pfk), a rate-limiting enzyme in glycolysis, on protein degradation pathways. Neuronal knockdown of Pfk was reported to lower ATP levels in brain neurons33. Pfk knockdown and milton knockdown decreased ATP to similar levels (Figure 3A-C). However, in contrast with milton knockdown, Pfk knockdown did not affect the levels of LC3-I, LC3-II or the LC3-II/LC3-I ratio (Figure 3D). Pfk knockdown decreased p62 level (Figure 3E), suggesting that autophagy is promoted. On the other hand, proteasome activity was decreased by Pfk knockdown (Figure 3F). These results suggest that the downregulation of axonal ATP upon depletion of axonal mitochondria decreases proteasome activity, but not autophagy.

ATP deprivation does not impair autophagy

(A-C) ATP levels in brain extracts of control and milton knockdown flies (A) and control and Pfk knockdown flies (B) and comparison of the effects of milton knockdown and Pfk knockdown on ATP levels (C). Flies were 14-day-old. Means ± SE, n = 3. (D, E) Western blotting of head extracts of flies with neuronal expression of control or Pfk RNAi. Blotting was performed with anti-LC3 (D) and anti-p62 (E) antibodies. For analyses of p62 levels, heads were extracted with 1% Triton X-100 or 2% SDS. Representative blots (left) and quantitation (right) are shown. Actin was used as a loading control. Means ± SE, n = 6 (LC3), n = 3 (p62). (F) Proteasome activity in head lysates of flies with neuronal expression of control or Pfk RNAi was measured by hydrolysis of Suc-LLVY-AMC. Means ± SE, n = 3. N.S., p > 0.05; *p < 0.05; **p < 0.01; ***p < 0.005 (Student’s t-test). Flies were at 14-day-old.

Proteome analysis suggests that depletion of axonal mitochondria causes disruption of autophagy and premature aging

To identify the pathways that mediate the decrease in autophagy in milton knockdown brains, we performed proteome analysis to systematically detect differentially expressed proteins upon neuronal knockdown of milton. We analyzed flies at 7- and 21-day-old, the age before autophagic defects are detected and the age just before the onset of neurodegeneration, respectively (Figure 4A). 1039 proteins were detected by liquid chromatography-tandem mass spectrometry (LC-MS/MS). Expression of 36 proteins was significantly increased (22 proteins) or decreased (14 proteins) by milton knockdown at 7-day-old (Figure 4B, Tables 1 and S1). At 21-day-old, the expression of 41 proteins (31 upregulated and ten downregulated proteins) was significantly altered in milton knockdown flies compared with control flies (Figure 4C, Tables 1 and S1). The “Interaction search” algorithm using KeyMolnet showed that proteins whose expression was significantly altered in the brains of milton knockdown flies at both 7- and 21-day-old were closely associated with the autophagic pathway (Table 2). Proteins involved in pathways characteristics of aging, such as the immune response (Transcriptional regulation by STAT), cancer (Transcriptional regulation by SMAD, Transcriptional regulation by myc), longevity (Transcriptional regulation by FOXO, Sirtuin signaling pathway), and stress responses (HSP90 signaling pathway, MAPK signaling pathway)34,35, were enriched in the proteome profiles of milton knockdown flies compared with those of control flies at 7-day-old (Table 2), suggesting that depletion of axonal mitochondria accelerates aging in the brain.

milton knockdown upregulates eIF2β in young flies

(A) Timing of proteome analysis and phenotypes observed in milton knockdown flies. (B, C) Scatter plots of the log2 abundance ratio (x-axis) against the -log10 p-value (y-axis) of proteins at 7-day-old (B) and 21-day-old (C). (D) eIF2 subunit protein levels from proteome analysis of milton knockdown flies compared to those of control flies. (E) Western blotting of head extracts of flies expressing control or milton RNAi in neurons with an anti-eIF2β antibody. Flies were 14-day-old. Representative blots (left) and quantitation (right) are shown. Tubulin was used as a loading control. Means ± SE, n = 6. (F) eIF2β mRNA levels quantified by qRT-PCR. Means ± SE, n = 4. (G) Western blotting of head extracts of wild-type flies with an anti-eIF2β antibody. Flies were 7-, 21-, 35-, 49- and 63-day-old. Representative blots (left) and quantitation (right) are shown. Tubulin was used as a loading control. Means ± SE, n = 3, *p < 0.05 (one-way analysis of variance (ANOVA) followed by Dunnett’s multiple comparison test).

Depletion of axonal mitochondria upregulates eIF2β and decreases phosphorylation of eIF2α

Differentially expressed proteins at 7-day-old flies may reflect alterations that are causal for autophagic defects. We noticed that the expression level of eIF2β was 2.465-fold higher in the brains of milton knockdown flies than in those of control flies (Figure 4B and D). Upregulation of eIF2β in the brains of milton knockdown flies was confirmed by western blotting. milton knockdown increased eIF2β protein levels more than twice (Figure 4E), but did not change the level of eIF2β mRNA (Figure 4F).

We also investigated age-dependent changes in eIF2β by western blotting of control flies at 7-, 21-, 35-, and 49-, and 63-day-old. eIF2β levels increased during aging until 49-day-old (Figure 4G). These results suggest that upregulation of eIF2β in milton knockdown fly brain reflects early an onset of age-dependent increase of eIF2β levels. eIF2β is a subunit of the eukaryotic initiation factor 2 (eIF2) complex, which is critical for translation initiation and the integrated stress response (ISR)4. eIF2 is a heterotrimer of α, β, and γ subunits, and eIF2α is phosphorylated during the ISR6. As for the other subunits of eIF2 complex, proteome analysis did not detect a significant difference in the protein levels of eIF2α and eIF2γ between milton knockdown and control flies at 7- and 21-day-old (Figure 4D). Western blotting of brain lysates showed that milton knockdown reduced eIF2α levels (Figure 5A), while p-eIF2α levels were not significantly affected (Figure 5B).

milton knockdown decreases phosphorylation of eIF2α

Earlier version: (A, B) Western blotting of head extracts with anti-eIF2α (A) and anti-p-eIF2α (B) antibodies. Flies were 14-day-old. Representative blots (left) and quantitation (right) are shown. Tubulin was used as a loading control. Means ± SE, n = 6. (C) A schematic representation of the axon (Lobe tips), the cell body region (Kenyon cells), and dendritic region (Calyxes) in the fly brain. Scale bars, 100µm. (D, E) Immunostaining with anti-eIF2α and anti-p-eIF2α antibodies. The mushroom body was identified by expression of mito-GFP. Scale bars, 20µm. The signal intensities of eIF2α and p-eIF2α in axons, dendrites, and cell bodies were quantified and are shown as ratios relative to the control. Means ± SE, n=12. N.S., p > 0.05; *p < 0.05; **p < 0.01; ***p < 0.005 (Student’s t-test).

Revised version: (A, B) Western blotting of head extracts with anti-eIF2α (A) and anti-p-eIF2α (B) antibodies. Flies were 14-day-old. Representative blots (left) and quantitation (right) are shown. Tubulin was used as a loading control. Means ± SE, n = 6. (C) A schematic representation of the axon (Lobe tips), the cell body region (Kenyon cells), and dendritic region (Calyxes) in the fly brain. Scale bars, 100µm. (D, E) Immunostaining with anti-eIF2α and anti-p-eIF2α antibodies. The mushroom body was identified by expression of mito-GFP. Scale bars, 20µm. The signal intensities of eIF2α and p-eIF2α in axons, dendrites, and cell bodies were quantified and are shown as ratios relative to the control. Means ± SE, nC=12. N.S., p > 0.05; *p < 0.05; **p < 0.01; ***p < 0.005 (Student’s t-test).

To analyze local changes of eIF2α and p-eIF2α, we carried out immunostaining. We focused on the mushroom body, where axons, dendrites, and cell bodies can be easily identified (Figure 5C). Both eIF2α and p-eIF2α were downregulated in the cell body (Kenyon cells) and dendritic (Calyxes) regions of the brains of milton knockdown flies (Figure 5D). In axons (lobe tips), milton knockdown did not affect eIF2α (Figure 5E, p = 0.271) but significantly downregulated p-eIF2α (Figure 5E). The ratio of p-eIF2α to eIF2α was lower in the axon but not in the soma or dendritic region. These results suggest that axonal distribution of mitochondria regulate the level of overall eIF2α protein and local p-eIF2α.

Depletion of axonal mitochondria suppressed global translation

Phosphorylation of eIF2α induces conformational changes in the eIF2 complex and inhibits global translation36. To analyze the effects of milton knockdown on translation, we performed polysome gradient centrifugation to examine the level of ribosome binding to mRNA. Since p-eIF2α was downregulated, we hypothesized that milton knockdown would enhance translation. However, unexpectedly, we found that milton knockdown significantly reduced the level of mRNAs associated with polysomes (Figure 6A and B). We also compared the level of translation between the brains of control and milton knockdown flies by assessing the incorporation of puromycin (Figure 6C). Puromycin incorporation was lower in the brains of milton knockdown flies than in those of control flies, while it was not statistically significant (Figure 6C, indicated by a bracket). These data suggest that the depletion of axonal mitochondria suppresses global translation.

milton knockdown suppressed global translation

(A) Representative polysome traces of head lysates of control and milton knockdown flies. (B) Quantitation of polysome fraction. The relative ratio of area under the curve (AUC) of polysome fractions (sedimentation 28%-50%). Means ± SE, n = 3. ***p < 0.005 (Student’s t-test) (C) Western blotting of head lysates of control and milton knockdown flies fed puromycin alone or puromycin and cycloheximide (CHX) with an anti-puromycin antibody. Flies were 14-day-old. Actin was used as a loading control. Representative blots (left) and quantitation (right) are shown. Means ± SE, n = 3. Student’s t-test.

eIF2β upregulation reduces the level of p-eIF2α, impairs autophagy, and decreases locomotor function

We were motivated to ask if eIF2β upregulation mediates autophagic defects caused by milton knockdown. If so, neuronal overexpression of eIF2β would also induce autophagy impairment. Neuronal overexpression of eIF2β increased LC3-II, while the LC3-II/LC3-I ratio was not significantly different (Figure 7A and B). Overexpression of eIF2β significantly increased the p62 level in the Triton X-100-soluble fraction (Figure 7C, 4-fold vs. control, p < 0.005 (1% Triton X-100)) but not in the SDS-soluble fraction (Figure 7C, 2-fold vs. control, p = 0.062 (2% SDS)), as observed in brains of milton knockdown flies (Figure 2B). These data suggest that neuronal overexpression of eIF2β accumulates autophagic substrates.

eIF2β upregulation impairs autophagy and decreases locomotor function

(A) eIF2β mRNA levels in head extracts of flies with UAS-eIF2β driven by elav-Gal4 (eIF2β OE) or UAS-GFP driven by elav-Gal4 (control) were quantified by qRT-PCR. Flies were 2-day-old. Means ± SE, n=4. (B, C) Western blotting of head extracts with anti-LC3 (B) and anti-p62 (C) antibodies. Flies were 14-day-old. Representative blots (left) and quantitation (right) are shown. Tubulin and actin were used as loading controls. Means ± SE, n = 3 (p62), n = 5 (LC3). (D, E) Western blotting of head extracts with anti-eIF2α (D) and anti-p-eIF2α (E) antibodies. Flies were 14-day-old. Representative blots (left) and quantitation (right) are shown. Tubulin was used as a loading control. Means ± SE, n = 6. (F) Climbing assay revealed early-onset of age-dependent locomotor defects in eIF2β-overexpressing flies. Means ± SE, n = 5. N.S., p > 0.05; ***p < 0.005 (Student’s t-test).

Since milton knockdown reduced the p-eIF2α level (Figure 5E), we asked whether an increase in eIF2β affects p-eIF2α. Neuronal overexpression of eIF2β did not affect the eIF2α level but significantly decreased the p-eIF2α level (Figure 7D and E).

Depletion of axonal mitochondria causes age-dependent decline in locomotor function24. We found that neuronal overexpression of eIF2β also caused locomotor dysfunction (Figure 7F). Locomotor functions were significantly impaired in those flies at 20-day-old and worsened further during aging (Figure 7F, compare 4-, 20-, and 30-day-old). We asked if eIF2β overexpression causes neurodegeneration, as depletion of axonal mitochondria in the photoreceptor neurons causes axon degeneration in an age-dependent manner24. eIF2β overexpression in photoreceptor neurons tends to increase neurodegeneration in aged flies, while it was not statistically significant (p>0.05, Figure S2).

These data indicate that an increase of eIF2β in neurons phenocopies depletion of axonal mitochondria, including suppression of autophagy and age-dependent locomotor dysfunction, and suggest that increase of eIF2β mediates these phenotypes downstream of loss of axonal mitochondria.

Lowering eIF2β rescues autophagic impairment and locomotor dysfunction induced by milton knockdown

Finally, we investigated whether suppression of eIF2β rescues autophagy impairment and locomotor dysfunction caused by neuronal knockdown of milton. Null mutants and flies with RNAi-mediated knockdown of eIF2β in neurons did not survive. Flies lacking one copy of the eIF2β gene survived without any gross abnormality, and the level of eIF2β mRNA in these flies was about 80% of that in control flies (Figure 8A). eIF2β heterozygosity did not affect the eIF2α and p-eIF2α levels (Figure S3A and B).

Lowering eIF2β rescues autophagic impairment and locomotor dysfunction induced by milton knockdown

(A) eIF2β mRNA levels with one disrupted copy of the eIF2β gene (eIF2βSAstopDsRed/+ (eIF2β -/+)). Head extracts of flies 2-3 day-old were analyzed by qRT-PCR. Means ± SE, n = 3. (B, C) Western blotting of head extracts of flies with neuronal expression of milton RNAi with or without eIF2β heterozygosity with anti-LC3 (B) and anti-p62 (C) antibodies. Flies were 14-day-old. Representative blots (left) and quantitation (right) are shown. Actin was used as a loading control. Means ± SE, n = 5 (LC3), n = 3 (p62). (D) The climbing ability of 20-day-old flies expressing milton RNAi with or without eIF2β heterozygosity. Means ± SE, n = 15. N.S., p > 0.05; *p < 0.05; ***p < 0.005 (Student’s t-test).

Neuronal knockdown of milton causes accumulation of autophagic substrate p62 in the Triton X-100-soluble fraction (Figure 2B), and we tested if lowering eIF2β ameliorates it. We found that eIF2β heterozygosity caused a mild increase in LC3-I levels and decreases in LC3-II levels, resulting in a significantly lower LC3-II/LC3-I ratio in milton knockdown flies (Figure 8B). eIF2β heterozygosity decreased the p62 level in the Triton X-100-soluble fraction in the brains of milton knockdown flies (Figure 8C). The p62 level in the SDS-soluble fraction, which is not sensitive to milton knockdown (Figure 2B), was not affected (Figure 8C). These results suggest that suppression of eIF2β ameliorates the impairment of autophagy caused by milton knockdown.

eIF2β heterozygosity also rescued locomotor dysfunction induced by milton knockdown. milton knockdown flies with eIF2β heterozygosity exhibited better locomotor function than milton knockdown alone (Figure 8D). The milton mRNA level was not increased in these flies, indicating that the rescue effect in the eIF2β heterozygous background was not mediated by an increase in the milton mRNA level (Figure S3C). These data suggest that eIF2β upregulation mediates autophagy impairment and locomotor dysfunction caused by the depletion of axonal mitochondria.

Discussion

Earlier version: The depletion of axonal mitochondria and accumulation of abnormal proteins are both characteristics of aged brains37,38. Proteostasis perturbations trigger the formation of pathological aggregates and increase the risks of neurodegenerative diseases during aging. By using neuronal milton knockdown to deplete mitochondria from the axon, we provide evidence that loss of axonal mitochondria drives age-related proteostasis collapse via eIF2β (Figure 9). We observed declines in autophagy-mediated degradation of less-aggregated proteins and proteasome activity in milton knockdown flies (Figure 2). Accumulation of ubiquitinated proteins and changes in age-related pathways started prematurely in milton knockdown flies (Figure 1 and Table 2). milton knockdown increased eIF2β and lowered eIF2α phosphorylation in young fly brain (Figure 4 and 5). Overexpression of eIF2β phenocopied the effects of milton knockdown, including reduced autophagy and accelerated age-related locomotor defects (Figure 7). Furthermore, lowering eIF2β levels suppressed the impairment of autophagy and locomotor dysfunction induced by milton knockdown (Figure 8). From these results, we propose that upregulation of eIF2β downstream of depletion of axonal mitochondria drives age-dependent collapse of proteostasis (Figure 9). Our results suggest that mitochondrial distribution and eIF2β are part of the mechanisms constituting proteostasis. milton knockdown causes loss of mitochondria in the axon and accumulation of mitochondria in the soma. Thus, the detrimental effects may be mediated by the accumulation of mitochondria. However, degeneration induced by milton knockdown is prominent in the axon and not detected in the cell body24. Furthermore, abnormal protein accumulation was observed in the axon (Figure 1), and p-eIF2α/eIF2α was decreased in the neurites but not in the soma (Figure 5), suggesting that proteostasis defects studied in this work are caused by depletion of mitochondria rather than accumulation of mitochondria. Further analyses to dissect the effects of milton knockdown on proteostasis and translation in the cell body and axon by experiments with spatial resolution would be needed.

Revised version: The depletion of axonal mitochondria and accumulation of abnormal proteins are both characteristics of aged brains37,38. Proteostasis perturbations trigger the formation of pathological aggregates and increase the risks of neurodegenerative diseases during aging. By using neuronal milton knockdown to deplete mitochondria from the axon, we provide evidence that loss of axonal mitochondria drives age-related proteostasis collapse via eIF2β (Figure 9). We observed declines in autophagy-mediated degradation of less-aggregated proteins and proteasome activity in milton knockdown flies (Figure 2). Accumulation of ubiquitinated proteins and changes in age-related pathways started prematurely in milton knockdown flies (Figure 1 and Table 2). milton knockdown increased eIF2β and lowered eIF2α phosphorylation in young fly brain (Figure 4 and 5). Overexpression of eIF2β phenocopied the effects of milton knockdown, including reduced autophagy and accelerated age-related locomotor defects (Figure 7). Furthermore, lowering eIF2β levels suppressed the impairment of autophagy and locomotor dysfunction induced by milton knockdown (Figure 8). From these results, we propose that upregulation of eIF2β downstream of depletion of axonal mitochondria drives age-dependent collapse of proteostasis (Figure 9). Our results suggest that mitochondrial distribution and eIF2β are part of the mechanisms constituting proteostasis.

Graphical abstract: the mitochondria-eIF2β axis in the axon maintains neuronal proteostasis during aging

Added in the revised version: milton knockdown causes loss of mitochondria in the axon and accumulation of mitochondria in the soma. Thus, the detrimental effects may be mediated by the accumulation of mitochondria. However, degeneration induced by milton knockdown is prominent in the axon and not detected in the cell body24. Furthermore, abnormal protein accumulation was observed in the axon (Figure 1), and p-eIF2α/eIF2α was decreased in the neurites but not in the soma (Figure 5), suggesting that proteostasis defects studied in this work are caused by depletion of mitochondria rather than accumulation of mitochondria. Further analyses to dissect the effects of milton knockdown on proteostasis and translation in the cell body and axon by experiments with spatial resolution would be needed.

Our results suggest that the loss of axonal mitochondria is an event upstream of proteostasis collapse during aging. The number of puncta of ubiquitinated proteins was higher in milton knockdown at 14-day-old, but there was no significant difference at 30-day-old (Figure 1). Proteome analyses also showed that age-related pathways, such as immune responses, are enhanced in young flies with milton knockdown (Table 2). We also found that eIF2β protein levels increase in an age-dependent manner until 49-day-old and reduces after that (Figure 4G). In the brains with neuronal knockdown of milton, eIF2β levels were higher at 7-day-old than those in control and lower at the 21-day-old (Figure 4D and Supplementary table). These results suggest that milton knockdown is likely accelerating age-dependent changes rather than increasing their magnitude. Disruption of proteostasis is expected to contribute neurodegeneration38, and it would be interesting to analyze the sequence of protein accumulation and axonal degeneration in milton knockdown (24,29 and Figure 1) in detail with higher time resolution.

Earlier version: Our results revealed that eIF2β regulates autophagy and maintains proteostasis during aging. eIF2β is a component of eIF2, which meditates translational regulation and ISR initiation. When ISR is activated, phosphorylated eIF2α suppresses global translation and induces translation of ATF4, which mediates transcription of autophagy-related genes39,40. Since ISR can positively regulate autophagy, we suspected that suppression of ISR underlies a reduction in autophagic protein degradation. We found neuronal knockdown of milton reduced phosphorylated eIF2α, suggesting that ISR is reduced (Figure 5). However, we also found that global translation was reduced (Figure 6). Increased levels of eIF2β might disrupt the eIF2 complex or alter its functions. The stoichiometric mismatch caused by an imbalance of eIF2 components may inhibit ISR induction. Supporting this model, we found that eIF2β upregulation reduced the levels of p-eIF2α (Figure 7). It is also possible that eIF2β mediates autophagy defects via mechanisms independent of ISR since eIF2β has functions independent of eIF241,42. For example, suppression of eIF2β has been reported to slow down cancer cell growth41. In developing neurons, eIF2β can directly interact with the translational repressor Kra to regulate midline axon guidance42. However, to our knowledge, the roles of eIF2β in aging have not been reported. Our results revealed a novel function of eIF2β to maintain proteostasis during aging.

Revised version: Our results revealed that eIF2β regulates autophagy and maintains proteostasis during aging. eIF2β is a component of eIF2, which meditates translational regulation and ISR initiation. When ISR is activated, phosphorylated eIF2α suppresses global translation and induces translation of ATF4, which mediates transcription of autophagy-related genes39,40. Since ISR can positively regulate autophagy, we suspected that suppression of ISR underlies a reduction in autophagic protein degradation. We found neuronal knockdown of milton reduced phosphorylated eIF2α, suggesting that ISR is reduced (Figure 5). However, we also found that global translation was reduced (Figure 6). Increased levels of eIF2β might disrupt the eIF2 complex or alter its functions. The stoichiometric mismatch caused by an imbalance of eIF2 components may inhibit ISR induction. Supporting this model, we found that eIF2β upregulation reduced the levels of p-eIF2α (Figure 7). It is also possible that eIF2β mediates autophagy defects via mechanisms independent of ISR since eIF2β has functions independent of eIF241,42. For example, suppression of eIF2β has been reported to slow down cancer cell growth41. In developing neurons, eIF2β can directly interact with the translational repressor Kra to regulate midline axon guidance42. Our results also suggest that milton knockdown and overexpression of eIF2β affect autophagy via increased LC3-I abundance (Figures 2 and 7), suggesting an unconventional mechanism of autophagy suppression. To our knowledge, the roles of eIF2β in aging and autophagy independent of ISR have not been reported. Our results revealed a novel function of eIF2β to maintain proteostasis during aging, while further investigation is required to elucidate underlying mechanisms.

How depletion of axonal mitochondria upregulates eIF2β is currently under investigation. A major mitochondrial function is ATP production, and depletion of axonal mitochondria downregulates ATP in axons33. However, we found that ATP deprivation did not always suppress autophagy (Figure 3), suggesting it is unlikely to be involved in the mechanisms that induce eIF2β upregulation. Mitochondria also serve as signaling hubs for translation and protein degradation. Mitochondrial proteins are regulated by co-translational protein quality control, and mitochondrial damage induces translational stalling of mitochondrial outer membrane-associated complex-I 30 kD subunit (C-I30) mRNA43. Additionally, the mitochondrial outer membrane ubiquitin ligase MITOL (also known as MARCHF5) ubiquitinates and regulates not only mitochondrial proteins such as Mfn244 but also microtubule-associated45 and endoplasmic reticulum46 proteins. These findings indicate that mitochondria serve as local signaling centers for proteostasis maintenance, and eIF2β levels may also be regulated by mechanisms related to mitochondria.

In conclusion, our results suggest that axonal mitochondria and eIF2β form an axis to maintain constitutive autophagy. Suppression of eIF2β rescued autophagic defects and neuronal dysfunction upon loss of axonal mitochondria. Since eIF2β is conserved across many species, including Drosophila and humans, our results suggest that eIF2β may be a possible therapeutic target for aging and diseases associated with mitochondrial mislocalization.

Materials and methods

Fly stocks and husbandry

Flies were maintained in standard cornmeal medium (10% glucose, 0.7% agar, 9% cornmeal, 4% yeast extract, 0.3% propionic acid, and 0.1% n-butyl p-hydroxybenzoate) at 25°C under light–dark cycles of 12:12 h. The flies were transferred to fresh food vials for every 2–3 days. UAS-milton RNAi (v41508) was from VDRC and outcrossed to [w1118] for five generations in our laboratory. Transgenic fly lines carrying UAS-luciferase RNAi control for milton RNAi) was reported previously24. GMR-gal4, Elav-gal4, UAS-Pfk RNAi (Bloomington stock center #36782), UAS-luciferase RNAi (Bloomington stock center #31603) (control for Pfk RNAi), UAS-GFP (used for control for UAS-eIF2β) and UAS-eIF2β (eIF2βEY08063, Bloomington stock center #17425) were from the Bloomington stock center. eIF2β heterozygous strain (PBac{SAstopDsRed} LL07719, DGRC#142114) was from KYOTO Drosophila Stock Center. UAS-mitoGFP was a kind gift from Drs. W. M. Saxton (University of California, Santa Cruz).

Immunohistochemistry and image acquisition

Fly brains were dissected in PBS and fixed for 45 minutes in formaldehyde (4% v/v in PBS) at room temperature. After incubation in PBST containing 0.1% Triton X-100 for 10 min three times, samples were incubated for 1h at room temperature in PBST containing 1% normal goat serum (Wako, #143-06561) and then incubated overnight with the primary antibody (anti-ubiquitin antibody Ubi-1 (Thermo Fisher #13-1600) (1:50), anti-eIF2α (abcam #ab26197) (1:50) and anti-p-eIF2α (Cell signaling #3398S) (1:50)) diluted in 1% NGS/PBST at 4 °C. Samples were then washed for 10 min with PBST including 0.1% Triton X-100 three times and incubated with the secondary antibody for overnight at 4° C. Brains were mounted in Vectashield (Vectorlab Cat#H-1100) and analyzed under a confocal microscope (Nikon). Quantitative analysis was performed using ImageJ (National Institutes of Health) with maximum projection images derived from Z-stack images acquired with same settings. Puncta were identified with mean intensity and area using ImageJ. For eIF2α and p- eIF2α immunostaining, the mushroom body was detected by mitoGFP expression.

Electron microscopy

Proboscis was removed from decapitated heads, which were then incubated in primary fixative solution (2.5% glutaraldehyde and 2% paraformaldehyde in 0.1 M sodium cacodylate buffer) at R.T. for 2 hours. After washing heads with 3% sucrose in 0.1 M sodium cacodylate buffer, fly heads were post-fixed for 1 hour in secondary fixation (1% osmium tetroxide in 0.1 M sodium cacodylate buffer) on ice. After washing with H2O, heads were dehydrated with ethanol and infiltrated with propylene oxide and Epon mixture (TAAB and Nissin EM) for 3 hours. After infiltration, specimens were embedded with an Epon mixture at 70°C for 2-3 days. Thin sections (70 nm) of laminas were collected on copper grids. The sections were stained with 5% uranyl acetate in 50% ethanol and Reynolds’ lead citrate solution. Electron micrographs were obtained with a CCD Camera mounted on a JEM-1400 plus electron microscope (Jeol Ltd.). Quantitation was performed using ImageJ (National Institutes of Health).

SDS–PAGE and immunoblotting

Western blotting was performed as reported previously24. Briefly, heads of 10-20 Drosophila were homogenized with SDS-Tris-Glycine sample buffer (0.312M Tris, 5% SDS, 8% glycerol, 0.0625% BPB, 10% β-mercaptoethanol, 10μg/mL leupeptin, 0.4 μM Pefabloc, 10mM β-glycerphosphate, 10mM NaF) and after boiling at 95°C for 2 minutes, it was centrifuged at 13,200 rpm, and the supernatant was used as a sample. For p62 western blot, fly heads were homogenized with 1% PBST and after centrifugation at 13,200 rpm, the supernatant was mixed 1:1 SDS-Tris-Glycine sample buffer, and boiled at 95 °C for 2 minutes. The pellet was dissolved with 2% SDS in PBS, then centrifuged again at 13,200 rpm. The supernatant was mixed 1:1 SDS-Tris-Glycine sample buffer and then boiled at 95 °C for 2 minutes. SDS–PAGE for western blotting was performed using 15%(w/v) (LC3), 7.5%(w/v) (p62), 10% (w/v) (eIF2α, β, and p-eIF2α) polyacrylamide gels. After electrophoresis, they were transferred to PVDF membrane (Merck Millipore) using a transfer device (BIO-RAD). After transfer, the membrane was blocked with 5% skim milk/TBST (50 mM Tris (pH 7.5), 0.15 M NaCl, 0.05% Tween20) for 1 hour and incubated with primary antibody listed below overnight at 4 °C. Membranes were rinsed twice with TBST containing 0.65M NaCl and once with TBST containing 0.15M NaCl. After incubation with the secondary antibody at room temperature for 1 hour, membranes were rinsed twice with TBST containing 0.65M NaCl and once with TBST containing 0.15M NaCl. After incubation with Immobilon Western Chemiluminescent HRP Substrate (Merck Millipore), chemiluminescent signals were detected with Fusion FX (Vilber). Experiments were repeated at least 3 times with independent cohorts of flies.

Primary antibodies: anti-LC3 antibody Atg8 (Merck Millipore #ABC975) (1:1000), anti-p62 antibody Ref2P (abcam #ab178440) (1:750), anti-eIF2β antibody (1:1500), anti-eIF2α antibody (abcam #ab26197) (1:1000), anti-p-eIF2α antibody (Cell signaling #3398S) (1:2000), anti-actin antibody (SIGMA A2066) (1:3000) and anti-β tubulin antibody (Sigma #T9026) (1:800,000). Polyclonal anti-eIF2β antibody was raised against a synthetic peptide (CGLEDDTKKEDPQDEA) corresponding to the C-terminal residues 29-43 of Drosophila eIF2β (1:1500).

Secondary antibodies: Peroxidase-conjugated goat anti-mouse IgG antibody (Dako #P0447) (1:2000), peroxidase-conjugated pig anti-rabbit IgG antibody (Dako #P0399) (1:2000)

Proteasome assay

Heads from ten flies were homogenized in 150 µl of buffer B (25 mM Tris-HCl [pH 7.5], 2 mM ATP, 5 mM MgCl2, and 1 mM dithiothreitol). Proteasome peptidase activity in the lysates was measured with a synthetic peptide substrate, succinyl-Leu-Leu-Val-Tyr-7-amino-4-methyl-coumarin (Suc-LLVY-AMC) (Cayman). Luminescence was measured on a multimode plate reader 2300 Enspire (PerkinElmer). Experiments were repeated at least 3 times with independent cohorts of flies.

ATP assay

Earlier version: Heads from the 10 flies were homogenized in 50 μl of 6 M guanidine-HCl in extraction buffer (100 mM Tris and 4 mM EDTA, pH 7.8) to inhibit ATPases. Samples were boiled for 5 min and centrifuged. The supernatant was diluted 4% with extraction buffer and mixed with a reaction solution (ATP Determination kit, Invitrogen). Luminescence was measured on a multimode plate reader 2300 Enspire (PerkinElmer). The relative ATP levels were calculated by dividing the luminescence by the total protein concentration, which was determined by the Bradford method. Experiments were repeated at least 3 times with independent cohorts of flies.

Revised version: Heads from the 10 flies were homogenized in 50Cμl of 6CM guanidine-HCl in extraction buffer (100CmM Tris and 4CmM EDTA, pH 7.8) to inhibit ATPases. Samples were boiled for 5Cmin and centrifuged. The supernatant was diluted 4% with extraction buffer and mixed with a reaction solution (ATP Determination kit, Invitrogen). Luminescence was measured on a multimode plate reader 2300 Enspire (PerkinElmer). The relative ATP levels were calculated by dividing the luminescence by the total protein concentration, which was determined by the Bradford method. Experiments were repeated at least 3 times with independent cohorts of flies.

Proteomic assay and pathway analysis

Sample preparation

Earlier version: Heads from the 35 flies were homogenized in 110 µl of extraction buffer (0.25% RapiGest SF, 50mM ammonium bicarbonate, 10mM dithiothreitol, 10μg/mL leupeptin, 0.4 μM Pefabloc, 10mM β-glycerphosphate, 10mM NaF). Homogenized samples were centrifuged and boiled for 5 min. After quantification of the protein concentration using a Pierce® 660 nm Protein Assay (Thermo Fisher Scientific), 10 µg proteins from each sample were reduced using 5 mM tris (2-carboxyethyl) phosphine hydrochloride (TCEP-HCl; Thermo Fisher Scientific) at 60°C for 1 h, alkylated using 15 mM iodoacetamide (Fujifilm Wako Pure Chemical, Osaka, Japan) at room temperature for 30 min, and then digested using 1.5 µg Trypsin Gold (Mass Spectrometry Grade; Promega, Madison, WI, USA) at 37°C for 17 h. The digests were acidified by the addition of trifluoroacetic acid (TFA), incubated at 37°C for 30 min, and then centrifuged at 17,000 ×g for 10 min to remove the RapiGest SF. The supernatants were collected and desalted using MonoSpin™ C18 (GL Sciences, Tokyo, Japan). The resulting eluates were concentrated in vacuo, dissolved in 2% MeCN containing 0.1% formic acid (FA), and subjected to LC-MS/MS analysis.

Revised version: Heads from the 35 flies were homogenized in 110 µl of extraction buffer (0.25% RapiGest SF, 50mM ammonium bicarbonate, 10mM dithiothreitol, 10μg/mL leupeptin, 0.4 μM Pefabloc, 10mM β-glycerphosphate, 10mM NaF). Homogenized samples were centrifuged and boiled for 5Cmin. After quantification of the protein concentration using a Pierce® 660 nm Protein Assay (Thermo Fisher Scientific), 10 µg proteins from each sample were reduced using 5 mM tris (2-carboxyethyl) phosphine hydrochloride (TCEP-HCl; Thermo Fisher Scientific) at 60°C for 1 h, alkylated using 15 mM iodoacetamide (Fujifilm Wako Pure Chemical, Osaka, Japan) at room temperature for 30 min, and then digested using 1.5 µg Trypsin Gold (Mass Spectrometry Grade; Promega, Madison, WI, USA) at 37°C for 17 h. The digests were acidified by the addition of trifluoroacetic acid (TFA), incubated at 37°C for 30 min, and then centrifuged at 17,000 ×g for 10 min to remove the RapiGest SF. The supernatants were collected and desalted using MonoSpin™ C18 (GL Sciences, Tokyo, Japan). The resulting eluates were concentrated in vacuo, dissolved in 2% MeCN containing 0.1% formic acid (FA), and subjected to LC-MS/MS analysis.

LC-MS/MS analysis and database search

LC-MS/MS analyses were performed on an Ultimate 3000 RSLCnano system (Thermo Fisher Scientific) coupled to a Q Exactive hybrid quadrupole-Orbitrap mass spectrometer (Thermo Fisher Scientific) equipped with a nano electron spray ionization (ESI) source. The LC system was equipped with a trap column (C18 PepMap 100, 0.3 × 5 mm, 5 µm, Thermo Fisher Scientific) and an analytical column (NTCC-360/75-3-125, Nikkyo Technos, Tokyo, Japan). Peptide separation was performed using a 90-min gradient of water/0.1% FA (mobile phase A) and MeCN/0.1% FA (mobile phase B) at a flow rate of 300 nL/min. Elution was performed as follows: 0–3 min, 2% B; 3–93 min, 2%–40% B; 93–95 min, 40%–95% B; 95–105 min, 95% B; 105–107 min, 95%–2% B; and 107–120 min, 2% B. The mass spectrometer was operated in data-dependent acquisition mode. The MS parameters were as follows: spray voltage, 2.0 kV; capillary temperature, 275°C; S-lens RF level, 50; scan type, full MS; scan range, m/z 350–1500; resolution, 70,000; polarity, positive; automatic gain control target, 3 × 106; and maximum injection time, 100 msec. The MS/MS parameters were as follows: resolution, 17,500; automatic gain control target, 1 × 105; maximum injection time, 60 msec; normalized collision energy (NCE), 27; dynamic exclusion, 15 sec; loop count, 10; isolation window, 1.6 m/z; charge exclusion: unassigned, 1 and ≥8; and injection volume, 1 µL (containing 0.5 µg protein). Measurements were made in duplicate for each sample.

The identification of proteins and label-free quantification (LFQ) of the detected peptides were performed using Proteome Discoverer software ver. 2.4 (Thermo Fisher Scientific). The analytical parameters used for the database search were as follows: parent mass error tolerance, 10.0 ppm; fragment mass error tolerance, 0.02 Da; search engine, sequest HT; protein database, Drosophila melanogaster (Fruit fly: SwissProt Tax ID=7227); enzyme name, trypsin (full); maximum number of missed cleavages, 2; dynamic modification, oxidation (methionine), phosphorylation (serine, threonine, tyrosine), acetyl (lysine), GG (lysine); N-terminal modification, Met-loss (methionine), and Met-loss+acetyl (methionine); static modification, carbamidomethylation (cysteine) and FDR confidence, High < 0.01, 0.01 ≤ Medium < 0.05, 0.05 ≤ Low. The parameters for LFQ were as follows: precursor abundance, based on area; and normalization mode, total peptide amount.

The abundance ratio of milton RNAi to control RNAi at 7- or 21-day-old was calculated. We considered proteins with an abundance ratio of ≥2.0 or ≤ 0.5 and an ANOVA P-value of < 0.05 based on volcano plots to be differentially expressed of milton RNAi. To extract molecular networks biologically relevant to the proteins that are differentially expressed in milton RNAi, pathway analysis was performed using KeyMolnet (KM Data Inc., Tokyo, Japan).

RNA extraction and quantitative real-time PCR analysis

Heads from more than 25 flies were mechanically isolated, and total RNA was extracted using ISOGEN (NipponGene) followed by reverse-transcription using PrimeScript RT reagent kit (Takara). The resulting cDNA was used as a template for PCR with THUNDERBIRD SYBR qPCR mix (TOYOBO) on a Thermal Cycler Dice real-time system TP800 (Takara). Expression of genes of interest was standardized relative to rp49. Relative expression values were determined by the ΔΔCT method (Livak and Schmittgen, 2001). Experiments were repeated three times, and a representative result was shown.

Primers were designed using DRSC FlyPrimerBank (Harvard Medical School). Primer sequences are shown below;

eIF2β for 5′-GGACGACGACAAGAGCGAAG -3′

eIF2β rev 5′-CGGTCGCATCACGAACTTTG -3′

milton for 5′-GGCTTCAGGGCCAGGTATCT-3′

milton rev 5′-GCCGAACTTGGCTGACTTTG-3′

Actin for 5′-TGCACCGCAAGTGCTTCTAA-3′

Actin rev 5′-TGCTGCACTCCAAACTTCCA-3′

Earlier version: rp49 for 5′-GCTAAGCTGTCGCACAAATG-3′ rp49 rev 5′- GTTCGATCCGTAACCGATGT-3′

Revised version: rp49 for 5′-GCTAAGCTGTCGCACAAATG-3′

Added in the revised version: rp49 rev 5′- GTTCGATCCGTAACCGATGT-3′

Polysome gradient centrifugation

30 heads were homogenized in 150 µl of lysis buffer (25 mM Tris pH 7.5, 50 mM MgCl2, 250 mM NaCl, 1 mM DTT, 0.5 mg/ml cycloheximide, 0.1 mg/ml heparin). The lysates were centrifuged at 13,200 rpm at 4 °C for 5 minutes, and the supernatant was collected. The samples containing 38µg of RNA was layered gently on top of a 10–50% w/w sucrose gradient (50 mM Tris pH 7.5, 50 mM MgCl2, 250 mM NaCl, 0.1 mg/ml heparin, 0.5 mg/ml cycloheximide in 5 ml polyallomer tube) and centrifuged at 37,000 rpm at 4°C for 150 minutes in a himac CP-NX ultracentrifuge using a P50AT rotor. Samples were fractionated from top to bottom, and absorbance at OD260 nm was analyzed by a Plate reader (EnSpire). Experiments were repeated at least 3 times with independent cohorts of flies.

Puromycin analysis

13-day-old flies were starved for 6 hr and fed 600 μM puromycin (Sigma) or 600 μM puromycin/35 mM cycloheximide (Sigma) in 5% sucrose solution for 20 hr. Incorporated puromycin was quantified by western blot with anti-puromycin antibody and normalized with actin. Experiments were repeated at least 3 times with independent cohorts of flies.

Histological analysis

Fly heads were fixed in Bouin’s fixative solution for 48 h at room temperature, incubated for 24 h in 50 mM Tris/150 mM NaCl, and embedded in paraffin. Serial sections (7 μm thickness) through the entire heads were stained with hematoxylin and eosin and examined by bright-field microscopy. Images of the sections that include the lamina were captured with Keyence microscope BZ-X700 (Keyence), and the vacuole area was measured using ImageJ (National Institutes of Health).

Climbing assay

The climbing assay was performed as previously described24. Flies were placed in an empty plastic vial (2.5 cm in diameter × 10 cm in length). The vial was gently tapped to knock the flies to the bottom, and the number of flies reached to the top, middle and bottom areas of the vials in 10 s was counted. Experiments were repeated 10 times, and the mean percentage of flies in each area and standard deviations were calculated. Experiments were repeated with independent cohorts more than three times, and a representative result was shown.

Statistics

The number of replicates, what n represents, precision measurements, and the meaning of error bars are indicated in Figure Legends. Data are shown as means ± SEM. For pairwise comparisons, Student’s t-test was performed with Microsoft Excel (Microsoft). For multiple comparisons, data were analyzed using one-way ANOVA with Tukey’s HSD multiple-comparisons test in the GraphPad Prism 6.0 software (GraphPad Software, Inc., La Jolla, CA). Results with a p-value of less than 0.05 were considered to be statistically significant.